---
  title: "second"
  description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: b824b67dc8   
```


First exercise


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: c3ab4b8ad6   
```


second exercise


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: a370122e02   
```





`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


***



```yaml
type: NormalExercise 
xp: 35 
key: 62f0adfd03   
```





`@instructions`
first

`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```



***



```yaml
type: NormalExercise 
xp: 35 
key: 4014c4b8cf   
```





`@instructions`
second

`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```



***



```yaml
type: NormalExercise 
xp: 30 
key: c7ef62e478   
```





`@instructions`
third

`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```



---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 85cabca573   
```


third exercise


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```

